```python
# Initialize Otter
import otter
grader = otter.Notebook("assignment_05.ipynb")
```

# Working with Branches & Forking

<!-- BEGIN QUESTION -->

<div class="alert alert-info">

## Submission instructions

<p>You receive marks for submitting your lab correctly, please follow these instructions:</p>

<ul>
  <!-- <li>Make at least three commits.</li>-->
  <li>Export to HTML via <code>File -> Save and export notebook as... -> HTML</code> and upload the HTML file to Canvas.</li>
    <ul>
      <li>Before submitting, make sure you restart the kernel and rerun all cells.</li>
    </ul>
  <li>Don't change any variable names that are given to you, don't move cells around, and don't include any code to install packages in the notebook.</li>
  </li>
</ul>
</div>

_Points:_ 1

<!-- END QUESTION -->

## Recap

In the past, we created `recipe` in Github.com and cloned (`git clone`) it to local PC.  If you recall, when you were typing `git status` into your terminal, the first line was always something like that `On branch main` and `Your branch is up to date with 'origin/main'`. So now, let's explore what it means in a bit more details.

Up until now our repository had only one branch (with default name`main`) with several commits coming after the other:

- `main` is typically the **default branch** name in git, and it represents the main development line of the repository. When you create a new repository, git will usually create a `main` branch by default. You can think of `main` as the local branch on your machine that you are currently working on. Please note that Github changed the `master` to be `main` recently as part of the company's effort to remove unnecessary references to slavery and replace them with more inclusive terms. Since this change is quite recent, you may notice that many reference materials/tutorials/website may still refer to this branch as `master`.

- `origin/main` is a branch that exists on the **remote** repository, also known as the **upstream** repository. When you clone a repository, you create a copy of the entire repository on your local machine. The `main` branch on your local machine reflects the `origin/main` branch on the remote repository. When you push changes from your local `main` branch to the remote repository, they will be added to the `origin/main` branch.

- `HEAD` is just a reference to the default branch in the repository. This can change if the default branch is renamed or deleted. If you are tracking the remote repository, then your local `HEAD` branch will point to `origin/HEAD`.

If you go to your local `recipe` folder from previous assignment and type `git log --oneline` you will see the following line at the 1st line:

`(HEAD -> main, origin/main, origin/HEAD)`

What does this line mean? It shows the relationships between the different branches and the `HEAD` pointer is a reference to the currently checked out branch or commit. It essentially tells git which branch or commit you are currently working on in your local or remote repository. The arrow symbol (`->`) is used to indicate that the `HEAD` pointer is pointing to the `main` branch. The `main` branch is the local branch you are currently working on. The comma-separated list of branches following the arrow represent the corresponding branches on the remote repository:

- `origin/main` is the name of the branch on the remote repository that corresponds to the `main` branch on your local repository.
- `origin/HEAD` is a reference to the default branch on the remote repository.

This line essentially shows you that the `HEAD` pointer is currently pointing to the main branch in your local repository, and that the `main` branch on your local repository corresponds to the `origin/main` branch on the remote repository. It also shows that `origin/HEAD` is being tracked by your local `HEAD` branch.

<div class="alert alert-info">
    
### Question 0
    
Imagine that a person is working on their git folder and check the output of the `git log --oneline`: 
    
```
a6f3f6d (HEAD -> feature, origin/feature) Add new feature
e5d2c1a Fix bug in existing feature
b3f1e2d Update documentation
f8c7a9e (origin/main) Initial commit
```
    
Which of the following sentences are TRUE?

A) The HEAD is pointing to the `feature` branch  
B) Three commits were done on `main` branch and one on `feature` branch  
C) One commit was done on `main` branch and one on `feature` branch  
D) The changes that were done on local `feature` branch is pushed to `origin/feature` remote branch  
E) The changes that were done on local `feature` branch is pushed to `origin/main` remote branch    
</div>

_Points:_ 1

```python
# Your answer should be a string with all the letters of the options you think are correct, e.g. 'AB'
q0 = 'ABD'
```

```python
grader.check("q0")
```

Output:
```
q0 results: All test cases passed!
```

## Creating a local branch

In simple terms, we use git branches to work on different features or changes to a project without affecting the main codebase or other people's work.Imagine that you're working on a "shark" project with several other people. You all want to try different variations of the "shark", add new features, and make changes. 

<img src="images/Picture2.png" width=70% >

If everyone worked on the same project at the same time, it could get confusing and difficult to keep track of who did what. It could also be risky if someone accidentally breaks the main code or deletes important work.That's where Git **branches** come in. When you create a new branch in Git (e.g. one called `shark`), your branch will not see any new changes made to the `main` branch; it will only be aware of the git history up until the point that it diverged from the `main` branch in the repository. This means that you can make changes, test new features, and experiment with new variations of the recipe on this branch without affecting the main codebase. Once you're satisfied with your changes and want to share them with others, you can merge your branch back into the main codebase. Git helps to track changes, resolve conflicts, and ensure that everyone's work is organized and synchronized.

----

**Step 1.** So let's create a new branch in our local `recipe` project to work on some feature improvements. Below is how our work will look like schematically:

<img src="images/Picture3.png" width=70% >


Let's crate a new branch called `fancy_readme`. If all worked, you will see a message in a terminal `Switched to a new branch 'fancy_readme'`.

<details><summary><b>Hint:</b></summary>
    
You can create a new branch by typing `git switch -c fancy_readme` in a terminal
    
</details>



<div class="alert alert-info">
    
### Question 1

Check the output of `git log --oneline` do you see any differences between the git history on `main` versus `fancy_readme` branches?
    
<details><summary><b>Hint:</b></summary>
    
You can use <code> git switch branch_name </code> command to switch between branches and checking the history in each branch 
    
</details>
<br>

A) `HEAD -> main, origin/main, origin/HEAD` changed to `HEAD -> fancy_readme, origin/main, origin/HEAD, main` when switching to `fancy_readme` branch   
B) all changes that were done on main branch is "checked out" into `fancy_readme` branch  
C) no changes been made, the output looks identical  
D) when we switch to `fancy_readme` branch, the git history is clear since no new commits have been made on a new branch 

</div>

_Points:_ 1

```python
# Your answer should be a string with all the letters of the options you think are correct, e.g. 'AB'
q1 = ...
```

```python
grader.check("q1")
```

----

**Step 2.**  Let's work on making our `README.md` file more fancy! Using the following guide to [basic markdown syntax](https://www.markdownguide.org/basic-syntax) (https://www.markdownguide.org/basic-syntax) do the following changes and make sure you make a **separate commit message for each of your changes**:

<details><summary><b>Hint:</b></summary>

You can use JupyterLab IDE to edit the `.md file`.
    
To make separate commit message means that after you complete each item in a list, stage and commit the changes providing a comprehensive commit message about the change 
    
</details>

1. Create level 1 and level 3 headers
2. Create ordered or unordered list
3. Make some words or sentences **bold** or *italicized*
4. Include the `recipe.png` or `recipe.jpg` into README.md

To preview the file in Jupyter IDE, right click on a file name, then select `Open With` in a menu and the select `Markdown Preview`. Please make sure that all the changes that you made are displayed as intended. 

<details><summary><b>Hint:</b></summary>

<img src="images/Picture4.png" width=80% > 
    
</details>

Once you done, make sure you don't have any unsaved changes. **DON'T push** anything to remote yet.



----

**Step 3.** 
So now we want to sync our local and remote repos! We usually use `git push` to send local changes to the remote. If you try it now, you will see that git wont push them, giving you the error: `fatal: The current branch fancy_readme has no upstream branch`. This error message  typically occurs when you try to push changes to a remote branch that your local branch is not tracking ( in our case there is no remote `fancy_readme` branch!). So git doesn't know where to push the changes because it doesn't have a remote branch to compare with the local branch.

To fix this error, you need to set up the tracking branch for your local branch. You can do this by running the following command:
`git push --set-upstream origin fancy_readme`. You will see that your git history now shows something similar to:


<img src="images/Picture7.png" width=50% >


After setting up the tracking branch, push your changes to the remote branch (you will no longer see the "no upstream branch" error). Schematically we can represent all our work we done so far in the following diagram:

<img src="images/Picture5.png" width=80% >

You can view your history using ` git log --all --graph --decorate --oneline` command in the terminal or using the History tab in Jupyter Lab.

<div class="alert alert-info">
    
### Question 2
    
Based on the output of `git log --all --graph --decorate --oneline` for a git repository:

```
* 8340b35 (HEAD -> test, origin/test) Add images to README
* 66b2f01 Update ingredients list
| * c1c9845 (main) Add recipe instructions to README
| * 2d70c8f Update README file
|/  
* 7e3c31d Initial commit
```

Which of the following statements are TRUE?

A) There are 2 branches: `main` branch and a `test` branch   
B) 2 commits were made on the `main` branch and 3 commits were made on the `test` branch  
C) changes on the `test` branch were pushed to the remote repository  
D) changes on the `main` branch were pushed to the remote repository,  
E) The `HEAD` pointer is pointing to the `test` branch.  
F) The changes on the `test` branch have been merged into the `main` branch  
    
</div>

_Points:_ 1

```python
# Your answer should be a string with all the letters of the options you think are correct, e.g. 'AB'
q2 = ...
```

```python
grader.check("q2")
```

----

**Step 4.** Merging the branches

Once we are happy with all the changes we did to our `README.md` file, let's bring all the changes we did on `fancy_readme` branch to our local `main` branch!

*Note: Usually we just delete the branch after we done working with it. So this step is just for you to practice merging branches or updating one branch with content of the other.*

First, what we need to do is to switch to `main branch` using `git switch main` and then run `git merge fancy_readme`

You can now see that all the changes are merged into `main branch` (you can check your `README.md` file too to make sure that all new changes are there). In the diagram below you can see both tags `main` and `fancy_readme` appearing on a same line:


<img src="images/Picture8.png" width=50% >

Also note, that we are now on `main` branch ( after running `git switch` command)

## Pull Request

In step 3, you have pushed the changes to your remote repository using the `git push`. However, all the new changes we did can only be seen on the `orign/fancy_readme` branch (i.e., remote `fancy_readme`) and not on our `main` branch (you can see it by checking `README` file on both branches):

<img src="images/Picture10.png" width=90% >    


Recall from step 4: we had a similar situation on our local PC: where we used the`git switch` and `git merge` commands to bring the changes from our `fancy_readme` to `main` branch.  So how to do the similar task on GitHub.com? We need to open a pull request:

<img src="images/Picture11.png" width=80% >    

-----

**Step 5.** Open a pull request.

Once the changes are pushed to your remote repository, go to the GitHub website and navigate to the `recipe` repository and switch to `fancy_readme` branch. 

<details><summary><b>Hint:</b></summary>
<img src="images/Picture10.png" width=90% >    
</details>


Click on the "Contribute" and select button "Open pull request".

<details><summary><b>Hint:</b></summary>
<img src="images/Picture12.png" width=90% >     
</details>

Once you press "Open pull request" you will see the following window: 

<img src="images/Picture13.png" width=90% >

At the top (1) you can see that we are merging `fancy_readme` branch into `main`. A checkmark "Able to merge" tells us that there is no conflicts and all the changes can be smoothly incorporated into the `main` branch. As always any changes you do on GitHub have to be accompanied with commit message (2). Before creating pull request (3) you can also:

- view your pull request (4) where you can review the changes that you made , 
- add a reviewer (5) if you work with collaborator or 
- add labels (6) - if you click on the wheel you can see a list of commonly used labels ( if you wish you can select any label that you think can be used to describe the changes we made)

Once you happy with all the changes,  submit the pull request by clicking the "Create pull request" button (3). Your pull request will then be sent for review.

In most Git workflows, pull requests are reviewed and accepted by the different person (not the one who did a pull request). These individuals are responsible for ensuring that the changes proposed in the pull request are aligned with the project goals, meet the quality standards, and do not introduce any issues or conflicts with the existing codebase.

Once you create a pull request you will see a page similar to the one below:

<img src="images/Picture14.png" width=90% >

As you can see from a page above you can accept the pull request by pressing a "Merge pull request" button or decline it by selecting the "Close pull request" option.Since you are working alone: you will be a person to approve pull request: 

You can go ahead and merge it! ( you will be asked to confirm a merge - you can leave the default message).

After the merge is done you have an option to delete `fancy_readme` branch. For the purpose of this exercise, don't delete the branch.

<!-- BEGIN QUESTION -->

<div class="alert alert-info">
    
### Question 3

Go back to the main page of your `recipe` repo: if pull request is done properly, you should be able to see the fancy version of your README file on a `main` branch. Then click on commits:

<img src="images/Picture15.png" width=90% >

Send us a URL link to your commit history (the link should look similar to this one: https://github.com/USERNAME/recipe/commits/main)

You will be marked on:
    
1. State of README from step 2
2. Quality of commit messages
3. Successfully merging branch

</div>

_Points:_ 3

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## Forking

In git, **forking** refers to making a copy of a repository on a user's GitHub account. The forked repository is a complete copy of the original repository, including all of its code, history, and branches. The forked repository is independent of the original repository and will not automatically get updates made to the original repository, so you would have to keep it in sync manually, which fortunately is straightforward to do via the GitHub interface.

Forking is usually done when a user wants to contribute to a project, but they don't have the permission to create branches and make changes inside the original repository. In this case, the user can fork the repository, make changes to the forked repository, and then submit a pull request to the original repository owner to merge their changes. As an example, this is how most big open source libraries, such as pandas, jupyter etc, receives contributions from.

<div class="alert alert-info">
    
### Question 4
    
Which of the following statements is true about branching and forking in git? Select one:

A) Branching and forking are the same thing.  
B) Branching does not create a copy of the repository but works off the history in the original repository, whereas forking creates a copy of the repository that is independent of the original repository.  
C) Forking creates a new branch in the repository that can be edited independently, while branching creates a copy of the repository that is linked to the original.  
D) Branching and forking both create a new branch in the repository that can be edited independently, but forking also creates a separate repository.    
    
</div>

_Points:_ 1

```python
# Your answer should be a string 
q4 = ...
```

```python
grader.check("q4")
```

<!-- BEGIN QUESTION -->

<div class="alert alert-info">
    
### Question 5
    
Using the steps, shown in this week video, fork this repository: [https://github.com/UBC-MDS/ds-toolbox-assignment5/](https://github.com/UBC-MDS/ds-toolbox-assignment5/)

**Step 1.** Clone your fork on your laptop.

**Step 2.** Create a new branch called `git-status`

**Step 3.** Create a new file called `git-status.md`

**Step 4.** In your own words, write a brief explanation of `git status` command in the `git-status.md` file. Use the title "Definition: git status" on the first line in your file. Specifically, explain what the `git status` command is useful for. How can you access the same information provided by this command if you were using the JupyterLab IDE instead of the terminal?

**Step 5.** Finally, save the changes, add them to the staging area, and commit and push your changes to your fork (the default push location).

**Step 6.** Go to the upstream repository (`UBC-MDS/ds-toolbox-assignment5`), click on Pull requests, and then click on Open pull request in the top right to create a pull request

After you completed these steps, send us a URL link to your pull request (the link should look similar to this one: `https://github.com/UBC-MDS/ds-toolbox-assignment5/pull/<NUMBER>`)
    
</div>

_Points:_ 3

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<div class="alert alert-danger" style="color:black">
    
**Restart and run all before submitting**
    
Before submitting,
don't forget to run all cells in your notebook
to make sure there are no errors
and so that the TAs can see the output of all the cells properly.
You can do this by clicking the ▶▶ button
or going to `Kernel -> Restart Kernel and Run All Cells...` in the menu.
This is not only important for this course,
but a good habit you should get into before ever committing a notebook to GitHub,
so that your collaborators can run it from top to bottom
without issues.
</div>

-----

Assignment 5 completed, Congratulations! 🥳

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Tests for q0 (points: 1)
```
>>> assert 'a' in q0.lower()
>>> assert 'b' in q0.lower()
>>> assert 'd' in q0.lower()
```
failure_message: You are missing one of the correct answers
```
>>> assert 'c' not in q0.lower()
>>> assert 'e' not in q0.lower()
```
failure_message: At least one of the options you have selected is not correct


Tests for q1 (points: 1)
```
>>> assert 'a' in q1.lower()
>>> assert 'b' in q1.lower()
```
failure_message: You are missing one of the correct answers
```
>>> assert 'c' not in q1.lower()
>>> assert 'd' not in q1.lower()
```
failure_message: At least one of the options you have selected is not correct


Tests for q2 (points: 1)
```
>>> assert 'a' in q2.lower()
>>> assert 'c' in q2.lower()
>>> assert 'e' in q2.lower()
```
failure_message: You are missing one of the correct answers
```
>>> assert 'b' not in q2.lower()
>>> assert 'd' not in q2.lower()
>>> assert 'f' not in q2.lower()
```
failure_message: At least one of the options you have selected is not correct


Tests for q4 (points: 1)
```
>>> assert 'b' == q4.lower()
```
failure_message: Try again! You answer is not correct
